```python
import numpy as np
```

Change the following variable to `True` if you want us to grade this challenge problem.

**IMPORTANT**: You can only get points for _one_ challenge problem per problem set. This means that even if you complete both challenge problems on this problem set, we will only count your best score between the two.

```python
SUBMIT_CHALLENGE = False
```

----

## Part A (1 point)

One of the places where production systems are still widely used is in making chatbots, which simulate human conversation. A chatbot production system consists of a set of rules that are triggered by keywords or phrases that appear in the input. For example, one rule might be "IF the input contains 'hello' THEN say 'hello'".

In this challenge problem, you will write a simple chatbot with at least 10 productions, capable of receiving sentences as input and producing responses. Your bot should be implemented as a function, `bot`, that takes in an input string and responds with an output string.

The following code runs a live conversation between you and the chatbot. The conversation is over when either of you responds with a string containing `bye` (e.g., "Goodbye.", "byebye", "goodbye").

```python
def chat(bot):
    ongoing = True
    while ongoing:
        user_input = input("You: ")
        bot_output = bot(user_input)
        print("Bot: {}".format(bot_output))
    
        if ("bye" in user_input.lower()) or ("bye" in bot_output):
            ongoing = False
```

Here is a simple bot that says hi.

```python
def hello_bot(user_input):
    return "Hello."
```

Talk to the bot a bit and then end the conversatoin.

```python
# Uncomment the following line.
#chat(hello_bot)
```

You can use any tools you'd like to create your bot. Here are a few ideas of functionality that you might include, with pointers to some of the Python modules that would be most useful in implementing that functionality. You'll have to remember to import any modules you want to use.
1. A memory that allows the bot to notice when the user repeats a sentence ( https://docs.python.org/3/tutorial/inputoutput.html#reading-and-writing-files).
1. The ability to answer math questions, maybe even with human-like timing ( https://docs.python.org/3/library/time.html#time.sleep).
1. Phone number recognition (https://docs.python.org/3/library/re.html).
1. A "super-duper-off" winner.
1. Some form of stochasticity (https://docs.python.org/3.0/library/random.html).

<div class="alert alert-success">
Edit the function `bot`, implementing a chat bot with at least 10 production rules.
</div>

<!-- grade_id: bot, points: 0, grade: false, solution: true -->
```python
def bot(user_input):
    """
    Hold a conversation, responding with a string containing "bye" to sign off.
    Your answer should not include spaces.

    Returns
    -------
    a string representing a sentence

    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

Here is a cell where you can test out your bot:

```python
# Uncomment the following line.
#chat(bot)
```

<!-- grade_id: check_for_submit, points: 1, grade: true, solution: false -->
```python
# Check whether you want us to grade the bot. 
# If this cell passes, we'll grade it manually; 
# otherwise we won't even look at it.
try:
    bot("bye")
except:
    pass

assert SUBMIT_CHALLENGE
```